evaluation_interval: 1m

rule_files:
  - api_server_availability_recording_rules.yaml

tests:
  # Start of a new test case. The interval and input series must be under this item.
  - name: Test KonfluxAPIServerHighErrorRate
    interval: 1m # Correct placement and spelling of interval
    input_series:
      - series: 'apiserver_request_total{code="503", job="kube-apiserver"}'
        values: '0+4x60'
      - series: 'apiserver_request_total{code="429", job="kube-apiserver"}'
        values: '0+1x60'
      - series: 'apiserver_request_total{code="200", job="kube-apiserver"}'
        values: '0+95x60'

    promql_expr_test:
      - expr: konflux_up
        eval_time: 1m
        exp_samples:
          - labels: 'konflux_up'
            value: 5
