evaluation_interval: 1m

rule_files:
  - kubearchive_availability_recording_rules.yaml

tests:
  - name: KubearchiveAvailabilityAllUpTest
    interval: 1m
    input_series:
      - series: "kube_deployment_status_replicas_available{namespace='product-kubearchive', deployment='d1'}"
        values: "3 3 3 3 3"
      - series: "kube_deployment_spec_replicas{namespace='product-kubearchive', deployment='d1'}"
        values: "3 3 3 3 3"

    promql_expr_test:
      - expr: konflux_up
        eval_time: 5m
        exp_samples:
          - labels: konflux_up{service='d1', check='replicas-available', namespace='product-kubearchive', deployment='d1'}
            value: 1

  - name: KubearchiveAvailabilitySomeUpTest
    interval: 1m
    input_series:
      - series: "kube_deployment_status_replicas_available{namespace='product-kubearchive', deployment='d1'}"
        values: "2 2 2 2 2"
      - series: "kube_deployment_spec_replicas{namespace='product-kubearchive', deployment='d1'}"
        values: "4 4 4 4 4"

    promql_expr_test:
      - expr: konflux_up
        eval_time: 5m
        exp_samples:
          - labels: konflux_up{service='d1', check='replicas-available', namespace='product-kubearchive', deployment='d1'}
            value: 0

  - name: KubearchiveAvailabilityMultipleDeploymentsTest
    interval: 1m
    input_series:
      # should be up (c1)
      - series: "kube_deployment_status_replicas_available{namespace='product-kubearchive', deployment='d1', source_cluster='c1'}"
        values: "1 1 1 1 1"
      - series: "kube_deployment_spec_replicas{namespace='product-kubearchive', deployment='d1', source_cluster='c1'}"
        values: "1 1 1 1 1"

      # should be down (c2)
      - series: "kube_deployment_status_replicas_available{namespace='product-kubearchive', deployment='d1', source_cluster='c2'}"
        values: "0 0 0 0 0"
      - series: "kube_deployment_spec_replicas{namespace='product-kubearchive', deployment='d1', source_cluster='c2'}"
        values: "1 1 1 1 1"

      # should be up (another deployment from c1)
      - series: "kube_deployment_status_replicas_available{namespace='product-kubearchive', deployment='d2', source_cluster='c1'}"
        values: "2 2 2 2 2"
      - series: "kube_deployment_spec_replicas{namespace='product-kubearchive', deployment='d2', source_cluster='c1'}"
        values: "1 1 1 1 1"
    promql_expr_test:
      - expr: konflux_up
        eval_time: 5m
        exp_samples:
          - labels: konflux_up{service='d1', check='replicas-available', namespace='product-kubearchive', deployment='d1', source_cluster='c1'}
            value: 1
          - labels: konflux_up{service='d1', check='replicas-available', namespace='product-kubearchive', deployment='d1', source_cluster='c2'}
            value: 0
          - labels: konflux_up{service='d2', check='replicas-available', namespace='product-kubearchive', deployment='d2', source_cluster='c1'}
            value: 1

  - name: KubearchiveAbsent
    interval: 1m
    input_series:
      - series: "kube_deployment_status_replicas_available{namespace='another-controller', deployment='d1'}"
        values: "1 1 1 1 1"
      - series: "kube_deployment_spec_replicas{namespace='another-controller', deployment='d1'}"
        values: "1 1 1 1 1"

    promql_expr_test:
      - expr: konflux_up
        eval_time: 5m
        exp_samples: []
