evaluation_interval: 1m

rule_files:
  - multi_platform_controller_availability_recording_rules.yaml

tests:
  - interval: 1m
    name: MPCExporterTest
    input_series:
      - series: "multi_platform_controller_available{service='multi-platform-controller-controller-manager-metrics-service', check='linux-arm64'}"
        values: "1 1 1 1 1"
      - series: "multi_platform_controller_available{service='multi-platform-controller-controller-manager-metrics-service', check='linux-ppc64le'}"
        values: "1 1 1 1 1"
      - series: "multi_platform_controller_available{service='multi-platform-controller-controller-manager-metrics-service', check='linux-amd64'}"
        values: "0 0 0 0 0"
      - series: "multi_platform_controller_available{service='multi-platform-controller-controller-manager-metrics-service', check='linux-s390x'}"
        values: "0 1 0 1 0"
    promql_expr_test:
      - expr: konflux_up
        eval_time: 5m
        exp_samples:
          - labels: konflux_up{service='multi-platform-controller', check='linux-arm64'}
            value: 1
          - labels: konflux_up{service='multi-platform-controller', check='linux-ppc64le'}
            value: 1
          - labels: konflux_up{service='multi-platform-controller', check='linux-amd64'}
            value: 0
          - labels: konflux_up{service='multi-platform-controller', check='linux-s390x'}
            value: 0
