evaluation_interval: 1m

rule_files:
  - workspaces_service_availability_recording_rules.yaml

tests:
  - interval: 1m
    name: WorkspaceExporterTest
    input_series:
      - series: "konflux_workspaces_available{service='workspaces-controller-manager-metrics'}"
        values: "1 1 1 1 1"
      - series: "konflux_workspaces_available{service='workspaces-controller-manager-metrics1'}"
        values: "0 0 0 0 0"
      - series: "konflux_workspaces_available{service='workspaces-controller-manager-metrics2'}"
        values: "0 1 0 1 0"
    promql_expr_test:
      - expr: konflux_up
        eval_time: 5m
        exp_samples:
          - labels: konflux_up{service='workspaces-controller-manager-metrics', check='kubesaw'}
            value: 1
          - labels: konflux_up{service='workspaces-controller-manager-metrics1', check='kubesaw'}
            value: 0
          - labels: konflux_up{service='workspaces-controller-manager-metrics2', check='kubesaw'}
            value: 0
