evaluation_interval: 1m

rule_files:
  - prometheus.gitops_argocd_app_reconcile_alerts.yaml

tests:

# Test that an alert is raised if only 90% of applications are reconciled within 4 seconds
  - interval: 1m
    input_series:
      - series: 'argocd_app_reconcile_bucket{namespace="gitops-service-argocd", le="+Inf"}'
        values: '1000+1000x35'
      - series: 'argocd_app_reconcile_bucket{namespace="gitops-service-argocd", le="4"}'
        values: '900+900x35'
    alert_rule_test:
      - eval_time: 35m
        alertname: GitOpsArgoCDAppReconcileAlert
        exp_alerts:
          - exp_labels:
              severity: warning
              slo: "true"
              namespace: gitops-service-argocd
            exp_annotations:
              summary: >-
                Less than 99% of Argo CD apps are reconciling within 4 seconds: 90%
              description: >-
                Argo CD in namespace gitops-service-argocd
                has failed to reconcile at least 99% of applications
                within 4 seconds over the past 30 minutes
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/gitops/argocd-app-reconcile.md

# Test that an alert is raised if only 98.9% of applications are reconciled within 4 seconds
  - interval: 1m
    input_series:
      - series: 'argocd_app_reconcile_bucket{namespace="gitops-service-argocd", le="+Inf"}'
        values: '1000+1000x35'
      - series: 'argocd_app_reconcile_bucket{namespace="gitops-service-argocd", le="4"}'
        values: '989+989x35'
    alert_rule_test:
      - eval_time: 35m
        alertname: GitOpsArgoCDAppReconcileAlert
        exp_alerts:
          - exp_labels:
              severity: warning
              slo: "true"
              namespace: gitops-service-argocd
            exp_annotations:
              summary: >-
                Less than 99% of Argo CD apps are reconciling within 4 seconds: 98.9%
              description: >-
                Argo CD in namespace gitops-service-argocd
                has failed to reconcile at least 99% of applications
                within 4 seconds over the past 30 minutes
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/gitops/argocd-app-reconcile.md

# Test that no alert is raised if 99% of applications are reconciled within 4 seconds
  - interval: 1m
    input_series:
      - series: 'argocd_app_reconcile_bucket{namespace="gitops-service-argocd", le="+Inf"}'
        values: '1000+1000x35'
      - series: 'argocd_app_reconcile_bucket{namespace="gitops-service-argocd", le="4"}'
        values: '990+990x35'
    alert_rule_test:
      - eval_time: 35m
        alertname: GitOpsArgoCDAppReconcileAlert
        exp_alerts: []

# Test that no alert is raised if 99.9% of applications are reconciled within 4 seconds
  - interval: 1m
    input_series:
      - series: 'argocd_app_reconcile_bucket{namespace="gitops-service-argocd", le="+Inf"}'
        values: '1000+1000x35'
      - series: 'argocd_app_reconcile_bucket{namespace="gitops-service-argocd", le="4"}'
        values: '990+990x35'
    alert_rule_test:
      - eval_time: 35m
        alertname: GitOpsArgoCDAppReconcileAlert
        exp_alerts: []

# Test that no alert is raised if there are no reconciliations
  - interval: 1m
    input_series:
      - series: 'argocd_app_reconcile_bucket{namespace="gitops-service-argocd", le="+Inf"}'
        values: '0x35'
      - series: 'argocd_app_reconcile_bucket{namespace="gitops-service-argocd", le="4"}'
        values: '0x35'
    alert_rule_test:
      - eval_time: 35m
        alertname: GitOpsArgoCDAppReconcileAlert
        exp_alerts: []
