evaluation_interval: 1m

rule_files:
  - 'prometheus.high_memory_usage_alert.yaml'

tests:
  - interval: 1m
    input_series:
      - series: "kube_pod_container_memory_usage_bytes{tenant_id='{{ $my_tenant_id }}', pod='{{ $my_pod}}', namespace='{{ $my_namespace }}'}"
        values: '1700000000 1800000000 1900000000 2000000000 2100000000'

      - series: "kube_pod_container_resource_limits_bytes{tenant_id='{{ $my_tenant_id }}', pod='{{ $my_pod}}', namespace='{{ $my_namespace }}'}"
        values: "2000000000 2000000000 2000000000 2000000000 2000000000"

    alert_rule_test:
      - eval_time: 5m
        alertname: high_memory_usage_alert
        exp_alerts:
          - exp_labels:
              severity: warning
              tenant_id: "{{ $my_tenant_id }}"
              pod: "{{ $my_pod }}"
              namespace: "{{ $my_namespace }}"
              slo: "true"
            exp_annotations:
              summary: "Pod {{ $my_pod}} in namespace {{ $my_namespace}} (Tenant: {{ $my_tenant_id }}) has exceeded it memory limit for more than 5 minutes"
              description: >
                "The pod {{ $my_pod}} (Tenant: {{ $my_tenant_id }}) has been consuming more than its defined limit in namespace {{ $my_namespace}} for at least 5 minutes. Consider investigating resource usage."
              alert_team_handle: <!subteam^S05Q1P4Q2TG>
              team: konflux-infra