evaluation_interval: 1m

rule_files:
  - prometheus.performance_alerts.yaml

tests:
  # Test for ETCD Alerts
  # - interval: 1m
  #   input_series:
  #     # Average Etcd FSync Latency is higher than 10ms, so it will be alerted.
  #     - series: 'etcd_disk_wal_fsync_duration_seconds_bucket{le="0.005"}'
  #       values: '0+10x10'  # Starts at 0 and increases by 10 for 10 steps
  #     - series: 'etcd_disk_wal_fsync_duration_seconds_bucket{le="0.01"}'
  #       values: '0+20x10'  # Starts at 0 and increases by 20 for 10 steps
  #     - series: 'etcd_disk_wal_fsync_duration_seconds_bucket{le="+Inf"}'
  #       values: '0+30x10'  # Starts at 0 and increases by 30 for 10 steps

  #   alert_rule_test:
  #     - eval_time: 10m
  #       alertname: EtcdFsyncLatency
  #       exp_alerts:
  #         - exp_labels:
  #             severity: warning


  # Node based Alerts
  # - interval: 1m
  #   input_series:
  #     # Node is running above 95% usage, so it will be alerted.

  #     # Simulating high CPU usage for instance1
  #     - series: 'node_cpu_seconds_total{mode="idle", instance="instance1"}'
  #       values: '1x20'  # Fast increase simulating high CPU usage
  #     # Simulating normal CPU usage for instance2
  #     - series: 'node_cpu_seconds_total{mode="idle", instance="instance2"}'
  #       values: '10x20'  # Slow increase simulating low CPU usage

  #   alert_rule_test:
  #     - eval_time: 10m
  #       alertname: NodeHighCPU
  #       exp_alerts:
  #         - exp_labels:
  #             severity: warning
  #             instance: instance1
  #           exp_annotations:
  #             summary: "Node High CPU Usage."
  #             description: "CPU Usage is above {{ $value }}% on node instance1"

  - interval: 1m
    input_series:
      # Node is running with 98% memory usage, so it will be alerted.
      - series: 'node_memory_MemTotal_bytes{instance="instance1"}'
        values: '100x20'  # Total Memory on Node
      - series: 'node_memory_MemAvailable_bytes{instance="instance1"}'
        values: '2x20'  # Free Memory on Node

    alert_rule_test:
      - eval_time: 10m
        alertname: NodeHighMemory
        exp_alerts:
          - exp_labels:
              severity: warning
              instance: instance1
            exp_annotations:
              summary: "Node High Memory Usage."
              description: "Memory Usage is above 98% on node instance1"

  - interval: 1m
    input_series:
      # Node is running under 90% memory usage, so it will not be alerted.
      - series: 'node_memory_MemTotal_bytes{instance="instance2"}'
        values: '100x20'  # Total Memory on Node
      - series: 'node_memory_MemAvailable_bytes{instance="instance2"}'
        values: '98x20'   # Free Memory on Node

    alert_rule_test:
      - eval_time: 10m
        alertname: NodeHighMemory
