evaluation_interval: 1m

rule_files:
  - 'prometheus.project_controller_up_alerts.yaml'

tests:
  - interval: 1m
    input_series:
      - series: 'konflux_up{namespace="project-controller", check="replicas-available", service="project-controller-controller-manager", source_cluster="c1"}'
        values: '0x15'
      - series: 'konflux_up{namespace="project-controller", check="replicas-available", service="project-controller-controller-manager", source_cluster="c2"}'
        values: '1x15'

    alert_rule_test:
      - eval_time: 12m
        alertname: ProjectControllerDown
        exp_alerts:
          - exp_labels:
              severity: critical
              check: replicas-available
              namespace: project-controller
              service: project-controller-controller-manager
              slo: "true"
              source_cluster: c1
            exp_annotations:
              summary: Project controller is down on cluster c1
              description: >
                Some of the replicas of project-controller-controller-manager are down
                on cluster c1
              alert_team_handle: <!subteam^S07NDQV6A4D>
              team: vanguard
              runbook_url: https://gitlab.cee.redhat.com/konflux/docs/sop/-/blob/main/vanguard/project-controller-down.md
